## Exercise 09.1 (checking data validity)

The Fibonacci series is valid only for $n \ge 0$. Add to the Fibonacci function in this notebook a check that raises an exception if $n < 0$. Try some invalid data cases to check that an exception is raised.

*Optional:* Use `pytest` to test that an exception *is* raised for some $n < 0$ cases.

<!-- grade_id: cell-982c3611466532a2, points: 0, grade: false, solution: true -->
```python
def f(n): 
    "Compute the nth Fibonacci number using recursion"
    if n <= 0:
        return 0
    elif n == 1:
        return 1
    else:
        return f(n - 1) + f(n - 2)
    
f(10)
```

Output:
```
55
```

<!-- grade_id: cell-9b97ea3e85d77f4a, points: 0, grade: true, solution: false -->
```python
## tests ##

# Perform some tests    
assert f(0) == 0
assert f(1) == 1
assert f(2) == 1
assert f(3) == 2
assert f(10) == 55
assert f(15) == 610

# Check that ValueError is raised for n < 0
import pytest
with pytest.raises(ValueError):
    f(-1)
with pytest.raises(ValueError):
    f(-2)
```

Output:
```
[... skipping hidden 1 frame]

File ~\AppData\Roaming\Python\Python313\site-packages\_pytest\outcomes.py:177, in fail(reason, pytrace)
    164 """Explicitly fail an executing test with the given message.
    165 
    166 :param reason:
   (...)    174     The exception that is raised.
    175 """
    176 __tracebackhide__ = True
--> 177 raise Failed(msg=reason, pytrace=pytrace)

Failed: DID NOT RAISE <class 'ValueError'>
```

## Exercise 09.2 (raising exceptions)

Modify your program from the bisection exercise in Activity 04 to raise an error if the maximum number of iterations is exceeded. Reduce the maximum allowed iterations to test that an exception is raised.

Add any other checks on the input data that you think are appropriate.

<!-- grade_id: cell-ca2c3ee62ad3b956, points: 0, grade: false, solution: true -->
```python
def my_f(x):
    """Evaluate polynomial function"""
    return x**5 / 10 + x**3 - 10 * x**2 + 4 * x + 7

def compute_root(f, x0, x1, tol, max_it):
    n=0
    error = tol + 1
    while error > tol and n < max_it:
        x_mid = (x0 + x1) / 2
        f_l = f(x0)
        f_mid = f(x_mid)

        if f_l * f_mid > 0:
            x0 = x_mid
        else:
            x1 = x_mid

        error = abs(f_mid)
        n += 1
    
    if n >= max_it:
        raise RuntimeError(
            f"Maximum number of iterations ({max_it}) exceeded before reaching tolerance {tol}."
        )
        
     
    return x_mid, f_mid, n
```

<!-- grade_id: cell-f6cf91ed5dd1736b, points: 0, grade: true, solution: false -->
```python
## tests ##

# Test with max_it = 30
x, f, num_it = compute_root(my_f, x0=0, x1=2, tol=1.0e-6, max_it=3)

# Test with max_it = 20
with pytest.raises(RuntimeError):
    x, f, num_it = compute_root(my_f, x0=0, x1=2, tol=1.0e-6, max_it=2)
```

Output:
```
Cell In[12], line 4
      1 ## tests ##
      2 
      3 # Test with max_it = 30
----> 4 x, f, num_it = compute_root(my_f, x0=0, x1=2, tol=1.0e-6, max_it=3)
      6 # Test with max_it = 20
      7 with pytest.raises(RuntimeError):

Cell In[9], line 22, in compute_root(f, x0, x1, tol, max_it)
     19     n += 1
     21 if n >= max_it:
---> 22     raise RuntimeError(
     23         f"Maximum number of iterations ({max_it}) exceeded before reaching tolerance {tol}."
     24     )
     27 return x_mid, f_mid, n

RuntimeError: Maximum number of iterations (3) exceeded before reaching tolerance 1e-06.
```